Functional specialization of mPFC-BLA and mPFC-NAc pathways in affective state representation

Author response:

The following is the authors’ response to the previous reviews

We sincerely thank the editors and reviewers for their careful evaluation and constructive feedback, which has helped us substantially improve the clarity and rigor of the manuscript. In the revised version, we have clarified the interpretation of the electrophysiological experiments, corrected the labeling of recorded signals as light evoked EPSCs, and removed statements implying differences in absolute synaptic strength. To address concerns about the interpretation of Fig. 7, we have added quantitative analyses of EPSC kinetics and revised the text to focus on synaptic response dynamics rather than amplitude differences. We have also removed analyses that could cause confusion and expanded the Methods section to provide additional experimental details, including the optogenetic stimulation configuration in slice recordings. Together, these revisions strengthen the interpretation of the electrophysiological results and improve the overall clarity and transparency of the study.

Public Reviews:

Reviewer #1 (Public review):

Weakness:

The authors focused primarily on female mice limiting generalizability and leaving the readers with questions about the impact of sex differences on their results. The tube test is used as a manipulation of the "emotional state" in several of the experiments. While the authors show the changes to corticosterone levels as a consequence of win/loss in the tube test, stronger claims might be made with comparisons to other gold standard stressors such as forced social defeat or social isolation.

We thank the reviewer for these thoughtful comments.

First, we acknowledge that the present study was conducted primarily in female mice, which may limit the generalizability of the findings. Female mice were selected to reduce variability associated with male aggression and housing-related stress, which can complicate behavioral assays such as social interaction and dominance testing. While focusing on a single sex allowed us to maintain experimental consistency across multiple behavioral paradigms, we agree that sex differences could influence the neural circuits underlying emotional and social behaviors. We have now added a statement in the Discussion acknowledging this limitation and noting that future studies will be necessary to determine whether similar circuit mechanisms operate in male mice.

Second, we appreciate the reviewer’s suggestion regarding the use of other stress paradigms. In this study, the tube test was used primarily to establish social dominance relationships between paired mice rather than as a classical stress-induction paradigm. Nevertheless, we observed measurable physiological changes associated with repeated win/loss outcomes, including alterations in corticosterone levels in brain lysates of loser mice after repeated tube-test competitions. Notably, repeated win/loss outcomes in the tube test were associated with significant increases in corticosterone levels in loser mice, indicating that the paradigm produced measurable physiological responses consistent with stress-related processes. These findings suggest that repeated social competition in this context can induce transient physiological and behavioral changes associated with social hierarchy. We agree that paradigms such as chronic social defeat stress or social isolation represent well-established models for inducing sustained stress responses. We have therefore revised the manuscript to clarify that the tube test in our study serves as a model of social competition and rank establishment rather than a canonical stress paradigm, and we highlight the comparison with other stress models as an important direction for future work.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

In relation to figure 7. Their response does not really clarify the issue:

(a) They argue that they are not making claims about synapse strength. However they still state "In the mPFC→NAc pathway, blue light stimulation evoked larger excitatory postsynaptic currents (EPSCs) in winner mice compared to losers (Fig. 7E). This suggests stronger synaptic transmission in winners' mPFC→NAc circuits. " They don't show this, they just show that normalized to some arbitrary value the responses of the earlier durations is higher or lower, which is very hard to interpret.

They argue in the rebuttal that the aim of this is to highlight response kinetics, but these are not quantified or discussed in any way.

We thank the reviewer for this helpful comment. We agree that the normalized input output curves shown in the original submission did not allow conclusions about absolute synaptic strength, and we also acknowledge that response kinetics were not previously quantified despite being mentioned in the rebuttal.

To address both concerns, we have revised Fig. 7 and added quantitative analyses of EPSC kinetics. Specifically, we measured the rise and decay slopes of light-evoked EPSCs recorded in postsynaptic neurons within the NAc and BLA of winner and loser mice. In the mPFC→BLA pathway, both the EPSC rise and decay slopes were significantly increased in loser mice compared with winners (rise slope: p = 0.0138; decay slope: p = 0.0392), suggesting enhanced synaptic responsiveness and faster charge transfer kinetics in BLA neurons of losers. In contrast, in the mPFC→NAc pathway, both mEPSC rise and decay slopes were not significantly different between groups.

These results provide a quantitative characterization of synaptic response dynamics and reveal pathway-specific differences in synaptic properties associated with social hierarchy. Importantly, this analysis does not rely on amplitude normalization and therefore allows a more interpretable comparison of synaptic response profiles between groups. We have updated Fig. 7 and the corresponding Results section to include these analyses.

(b) They still haven't labeled the responses correctly. The responses in figure 7 are not "voltage spikes" but light-evoked EPSCs.

We apologize for the incorrect terminology. All instances of “voltage spikes” have been corrected to “light-evoked EPSCs” in the figure legends and text.

(c) They argue that responses do not vary across experiments/slices because they use a constant viral injection volume targeted to the same co-ordinates and identical placement of the fiber and recording location. While I am sure they aim to do that, it is almost impossible to ensure that this was identical across experiments and that the degree of opsin labelling in their slices was the same (See for example Mao et al., 2011 PMID: 21982373 who pioneer the approach of using within slice comparisons to account for this). If I understand their explanation of their strategy correctly, the authors own rebuttal highlights this point, they seem to have needed to vary the LED duration by an order of magnitude (1-10ms) to ensure reliable responses across experiments, even for the same projection.

We thank the reviewer for raising this important point. We agree that it is not possible to ensure identical opsin expression or light delivery across experiments. We have revised the manuscript to explicitly acknowledge this limitation and clarify that normalization was used to mitigate, but not eliminate, inter-slice variability. We now avoid any interpretation that relies on absolute response amplitude across animals.

Regarding “LED duration variability (1-10 ms)”, we agree that the need to adjust stimulation duration reflects variability in effective opsin activation across slices. We now clarify this point in the Methods and Results and emphasize that stimulation parameters were optimized to reliably evoke responses rather than to equate absolute light input across experiments.

Importantly, our main conclusions do not rely on absolute EPSC amplitude comparisons. Instead, they are supported by analyses that are less sensitive to variability in opsin expression or light delivery, including EPSC kinetics (rise and decay slopes), paired-pulse ratio measurements, and AMPA/NMDA ratios. These complementary measures provide a more robust characterization of synaptic properties across conditions.

(d) Similarly in Fig S6 it is unclear what they are showing. The Y axis is still labeled in pA, yet they claim this is an action potential? Also this analysis is rather irrelevant to the data shown in figure 7 as the pathway between PFC and BLA/NAc is not preserved.

We thank the reviewer for pointing out the lack of clarity in Fig. S6. We agree that it does not directly inform the interpretation of Fig. 7 and may cause confusion. To improve the clarity and focus of the manuscript, we have therefore removed Fig. S6 from the revised manuscript. The removal of this supplementary figure does not affect the main conclusions of the study.

(e) It now also seems that these experiments were performed by placing a fiber optic into the slice to elicit responses. This should be detailed in the methods.

We thank the reviewer for noting this omission. We have added a detailed description of fiber-optic placement within the slice for optogenetic stimulation to the Methods section. Specifically, we clarify that blue light was delivered through a fiber optic positioned above the recorded slice to activate ChR2-expressing mPFC axon terminals within the BLA or NAc. The placement of the fiber relative to the recorded neurons and the stimulation parameters are now explicitly described in the revised Methods section.

Changes from the earlier version to the revised version

Abstract

Earlier version: Effective emotional processing, crucial for adaptive behavior, is mediated by the medial prefrontal cortex (mPFC) via connections to the basolateral amygdala (BLA) and nucleus accumbens (NAc), traditionally considered functionally similar in modulating reward and aversion responses. However, the functional specialization of the mPFC→BLA and mPFC→NAc pathways in representing affective states remains unclear. We found that while overall firing patterns appeared consistent across emotional states, deeper analysis revealed distinct variabilities. Specifically, mPFC→BLA neurons, especially “center-ON” neurons, exhibited heightened activity during behaviors classically associated with anxiety-like states, suggesting their involvement in aversive behavioral regulation. Conversely, mPFC→NAc neurons were more active during exploratory and approach-related behaviors, implicating them in the processing of positively valenced behavioral states. Notably, mPFC→NAc neurons showed significant pattern decorrelation during social interactions, suggesting a pivotal role in processing social preference. Additionally, chronic emotional states affected these pathways differently: positively valenced contexts enhanced mPFC→NAc activity, while negatively valenced conditions boosted mPFC→BLA activity. These findings challenge the assumed functional similarity and highlight distinct correlational patterns, suggesting potential, but not yet causally established, roles of mPFC→BLA and mPFC→NAc pathways in shaping emotional states.

Revised version: Effective emotional processing, crucial for adaptive behavior, is mediated by the medial prefrontal cortex (mPFC) via connections to the basolateral amygdala (BLA) and nucleus accumbens (NAc), traditionally considered functionally similar in modulating reward and aversion responses. However, the functional specialization of the mPFC→BLA and mPFC→NAc pathways in representing affective states remains unclear. We found that while overall firing patterns appeared consistent across emotional states, deeper analysis revealed distinct variabilities. Specifically, mPFC→BLA neurons, especially “center-ON” neurons, exhibited heightened activity during behaviors classically associated with anxiety-like states, suggesting their involvement in aversive behavioral regulation. Conversely, mPFC→NAc neurons were more active during exploratory and approach-related behaviors, implicating them in the processing of positively valenced behavioral states. Notably, mPFC→NAc neurons showed significant pattern decorrelation during social interactions, suggesting a pivotal role in processing social preference. Additionally, chronic emotional states affected these pathways differently: positively valenced contexts enhanced mPFC→NAc activity, while negatively valenced conditions boosted mPFC→BLA activity. Consistent with this behavioral divergence, repeated win/loss outcomes in the tube test were associated with elevated corticosterone levels in loser mice, indicating that repeated social competition produces measurable physiological responses linked to social hierarchy. These findings challenge the assumed functional similarity and highlight distinct correlational patterns, suggesting potential, but not yet causally established, roles of mPFC→BLA and mPFC→NAc pathways in shaping emotional states.

Introduction

Recent advancements in neuroscience have significantly advanced our understanding of the complex roles played by the medial prefrontal cortex (mPFC) pathways to the basolateral amygdala (BLA) and nucleus accumbens (NAc). These pathways are pivotal in emotional and motivational processing, mediating behaviors that involve assessing conflicts between reward and aversion, thus enabling the brain to adaptively respond to diverse environmental stimuli (1, 2).

Both the mPFC→BLA and mPFC→NAc pathways play key roles in shaping emotional responses and guiding decisions when animals face potential rewards or threats (3). These circuits adjust behavior when reward seeking must be suppressed in the presence of aversive cues, thereby maintaining a balance between risk and reward in emotionally charged contexts. The prelimbic (PL) subdivision of the mPFC sends excitatory projections to the BLA, a pathway causally linked to fear and anxiety: optogenetic activation of PL→BLA maintains fear memories and induces anxiety-like behaviors (4, 5). In contrast, the mPFC→NAc pathway is implicated in exploration, reward-seeking, and motivational drive, with optogenetic activation of PL→NAc projections biasing exploratory behavior and social approach (6).

The mPFC-BLA circuit also integrates reward-aversion conflicts. Ishikawa et al. showed that infralimbic (IL) projections to the BLA are required to suppress reward-seeking when paired with shock punishment (7), an effect likely mediated through BLA influence on the NAc shell (8). Inactivation of PL, BLA, or NAc shell disrupts avoidance (9, 10). PL projections to BLA and NAc have opposing effects on avoidance: PL-BLA activation facilitates avoidance, while PL-NAc activation inhibits it (11). During avoidance retrieval, PL neurons projecting to BLA are activated, while avoidance extinction activates both PL-NAc and IL-NAc projections (12). Finally, Kim et al. reported that activation of mPFC→NAc lateral shell neurons engaged by aversive stimuli can suppress reward seeking, underscoring the role of this pathway in weighing competing motivational demands (13).

Together, these studies highlight both shared and distinct roles of the mPFC→BLA and mPFC→NAc pathways in shaping emotional and motivational behaviors. While both circuits modulate responses to reward-threat conflict, they do so through partially divergent mechanisms, emphasizing the complexity of prefrontal control over affect and action. Despite substantial progress, however, many questions about the precise neural mechanisms within these pathways remain unresolved. Traditional approaches, including pharmacological manipulations, c-fos mapping, circuit tracing, and even pathway-specific optogenetics, have provided critical insights but lack the resolution to selectively manipulate emotion-related ensembles. Because these techniques often do not distinguish between specific cell types or functionally defined subpopulations, important circuit- and cell-level differences may remain obscured, limiting our understanding of how these pathways implement their distinct functions.

Building on this foundation, we employed in vivo imaging to examine the network dynamics of the mPFC→BLA and mPFC→NAc pathways during emotionally salient behaviors. Our goal was to directly compare how these circuits encode affective states, focusing on the similarities and differences in their neuronal activation patterns and ensemble dynamics. Specifically, we analyzed the temporal and spatial characteristics of population activity and their correlation with defined emotional behaviors. To probe these processes, we used a social competition paradigm to elicit affective states in mice, enabling us to capture pathway-specific plasticity in a naturalistic context. Through this approach, we uncovered distinct circuit dynamics and plasticity signatures in the two pathways, each aligning with behavioral states of different affective valence.

Results

Differential activity patterns of mPFC→NAc and mPFC→BLA neurons during anxiety-like and exploratory behavioral states

To distinguish the mPFC→BLA and mPFC→NAc neurons, we employed retrograde AAV viruses carrying genes encoding the Ca2+ indicator GCaMP6m, delivering them to the BLA and NAc of adult female wildtype mice at the age of ∼20 weeks, respectively (Fig. 1A). This allowed for the infection of mPFC neurons projecting to either of these regions, leading to GCaMP6 expression in the soma situated in the mPFC (Fig. 1B). In vivo Ca2+ imaging was conducted on mPFC→BLA and mPFC→NAc neurons using a portable miniaturized microscope (Fig. 1C). The activity of individual neurons was continuously monitored throughout the behavioral assays. We synchronized the Ca2+ imaging with the behavioral recordings, following the same neurons within each mouse across the different test sessions (see Methods).

The mPFC→NAc and mPFC→BLA neurons demonstrated distinct activity patterns across various emotional states.

(A) Schematic showing retrograde AAV-GCaMP6 injection in the NAc and BLA, respectively. In vivo calcium imaging was performed in the mPFC neurons expressing GCaMP6 via miniscope. (B) Top: An example of AAVretro-GCaMP6 injection and expression in the BLA and NAc, respectively. Bottom: An example of GCaMP6m expression in the mPFC with AAVretro-GCaMP6 injection in the NAc and an image of Ca2+ fluorescence recorded with miniaturized microscope. The 1-mm GRIN lens covered the full depth of the prelimbic cortex. Green, GCaMP6m. Scale bars, 200 mm. (C) Left: the field of view under a GRIN lens in one mouse with identified neurons numbered and colored. Right: fluorescence traces of example neurons marked in the above panel. (D). Schematic of the open field test (OFT) in an open 50 cm X 50 cm arena. The pink areas represent corners and the blue area represents the center. (E) Illustrations depicting typical mouse behaviors associated with different emotional states observed during the OFT: Center (exploration and higher anxiety), Corner (lower anxiety), sniffing (sensory exploration and vigilance, possibly indicating curiosity, cautious exploration, or heightened alertness), and grooming (stress relief or comfort). (F) The percentage of time that mice spent in the center, corners, sniffing and grooming during 10-min open field test. n = 15 mice for mPFC→BLA group; n = 11 mice for mPFC→NAc group. The time-allocation pattern across behaviors is similar for Amy and NAc cohorts. A two-way mixed-effects model (Region × Behavior) showed a strong main effect of Behavior (Corner » Center; Sniffing < Center), but no main effect of Region and no Region × Behavior interaction (all P > 0.1). (G) Heatmap and time-series plot of neuronal activity (ΔF/F) across the test duration. The heatmap above shows the fluctuation in activity level across neurons, while the plot below aligns the averaged fluctuations of these neurons with observed behaviors (Center, Corner, Sniffing, Grooming) within the arena. (H) Normalized transient rates of neuronal activity during different behaviors. The mixed-design ANOVA revealed no main effect of region (P = 0.14) and no significant Region × Behavior interaction (P > 0.3), indicating that mPFC→BLA and mPFC→NAc neurons exhibited similar behavioral modulation patterns. (I) PCA plots for mPFC→BLA and mPFC→NAc pathways illustrating the distribution of neuronal activity patterns for different observed behaviors (Center: blue, Corner: red, Sniffing: green, Grooming: purple). Each point represents an individual neuronal recording. Quantitative analysis confirmed that neural population activity during corner behavior was highly similar between mPFC→BLA and mPFC→NAc pathways (Pearson r = 0.64, p < 1 × 10⁻⁹), consistent with comparable encoding of this low-anxiety state. (J) Summary of the distances of neuronal activity clusters from the center of Corner behavior in the OFT for both pathways. distances were computed using PC1 and PC2 coordinates. ***P < 0.001, Mann-Whitney U test.

We first performed the open field test (OFT), a widely employed behavioral assay in rodent research for assessing exploratory and anxiety-like behaviors (14) (Fig. 1D). This test leverages the inherent conflict rodents experience between their desire to explore new environments and their instinctive aversion to exposed spaces. We monitored and recorded the mouse’s movements and actions over a 10-minute period, specifically tracking time spent in the center versus the corner, wall-sniffing, and grooming. These behaviors observed during the open field test, such as spending time in corners, the center area, grooming, and sniffing walls (Fig. 1E), offer valuable insights into the mice’s affectively salient behavioral states. Typically, mice in corners exhibit lower anxiety levels, suggesting a sense of security. Conversely, exposure to the center area may induce higher anxiety due to increased exposure and reduced safety compared to the arena’s edges or corners. Grooming behaviors can signal stress relief or comfort, while wall-sniffing reflects sensory exploration and vigilance, possibly indicating curiosity, cautious exploration, or heightened alertness (14, 15). Our findings demonstrate that mice consistently spend significant time in corners during the 10-minute session. Additionally, they allocate around 10% of the time to the center area or grooming, and approximately 5% of the time to wall-sniffing (Fig. 1F). To ensure accurate behavioral categorization, “sniffing” was defined as episodes in which the mouse’s nose was within ∼1 cm of the wall and moving in small scanning motions without locomotion for ≥1 s, as detected by the TopScan system. Sniffing events occurring within the corner zones were excluded from the corner category to maintain mutual exclusivity between the two behaviors. Accordingly, the corner time in Fig. 1F does not include sniffing episodes, allowing for a clear distinction between directed exploratory (sniffing) and stationary, anxiety-associated (corner) states.

To determine whether the pattern of neuronal activities in the mPFC→BLA and mPFC→NAc pathways differs in response to distinct affective behavioral states, we first estimated the Ca2+ transient rate. During our observation of mPFC→BLA and mPFC→NAc neuron activity (Fig. 1G) during the open field test, we determined that the collective averaged Calcium transient rate of the neurons in these two pathways did not exhibit discernible differences while the mice were situated in the center, corner, or during sniffing or grooming (Fig. 1H and Supplementary Fig. 1). To delve deeper, we conducted a Principal Component Analysis (PCA) to assess the population activity pattern linked to various behaviors, encompassing center or corner dwelling, grooming, and sniffing. The PCA plots in Fig. 1I reveal how neuronal activities associated with various behaviors differ between these pathways, with each point representing recorded neuronal activity during specific behaviors projected onto the first two principal components (PC1 and PC2). These components capture the most significant variances within the dataset. These two plots display distinct clustering of activity patterns with clear separations among the different behaviors, particularly between center and corner behaviors, as well as sniffing and grooming. To account for the apparent spread in the mPFC→BLA trajectories during sniffing, we examined the temporal progression of neural activity and found that early frames clustered tightly, whereas later frames diverged along PC1, reflecting transitions into other behaviors. Shortening the analysis window to 2-3 s post-onset reduced this spread, confirming that the variance observed in the 5 s window arises from within-bout behavioral transitions rather than overlap with corner activity or time-in-arena effects.

To quantify the similarity of behavioral representations between the two projection pathways, we performed a canonical correlation analysis (CCA) on the reduced PCA representations of averaged behavioral bouts. Corner behavior exhibited the smallest cross-pathway difference (mean absolute difference = 0.12, p = 0.00039 < 0.001), indicating highly overlapping representations between mPFC→BLA and mPFC→NAc neurons. In contrast, Center (0.47, p = 0.968), Sniffing (0.43, p = 0.0002), and Grooming (0.34, p = 0.0002) showed substantially larger differences. These findings confirm that corner activity is encoded most similarly across the two pathways, supporting its use as a stable, low-anxiety reference state for comparison.

Earlier version: We next quantified distances from Corner behavior to other behavioral states across both pathways (Fig. 1J). While both pathways exhibited similar neural patterns for Corner behavior, Center and Grooming states were markedly farther from Corner in the mPFC→BLA pathway, indicating more distinct neural encoding of anxiety-related versus non-anxiety behaviors. This differentiation was especially pronounced for Grooming, suggesting that the mPFC→BLA pathway plays a key role in distinguishing stress– and coping-related states. Conversely, the mPFC→NAc pathway showed stronger separation between Sniffing and Corner behaviors, highlighting its selective involvement in differentiating exploratory behaviors from those associated with anxiety (Fig. 1I, J).

Revised version: We next quantified distances from Corner behavior to other behavioral states across both pathways (Fig. 1J). Distances between behavioral clusters were calculated in the PC1-PC2 space as Euclidean distances from the centroid of the corner distribution, which served as a reference state because mice spent the majority of time in the corners. While both pathways exhibited similar neural patterns for Corner behavior, Center and Grooming states were markedly farther from Corner in the mPFC→BLA pathway, indicating more distinct neural encoding of anxiety-related versus non-anxiety behaviors. This differentiation was especially pronounced for Grooming, suggesting that the mPFC→BLA pathway plays a key role in distinguishing stress-and coping-related states. Conversely, the mPFC→NAc pathway showed stronger separation between Sniffing and Corner behaviors, highlighting its selective involvement in differentiating exploratory behaviors from those associated with anxiety (Fig. 1I, J).

To ensure that the observed distance relationships were not driven by random label structure or temporal correlations, we implemented permutation-based null controls. Behavioral labels were shuffled within each mouse and session, and circular time-shifts were applied to each neuron within bouts to preserve autocorrelation. The observed Corner distances and between-pathway differences exceeded the 95th percentile of both null distributions and remained significant after 10,000 permutations (p_perm < 0.01). These controls confirm that the reported cluster relationships reflect genuine structure in neural population activity. Together, these analyses underscore the distinct functional roles of the two mPFC projection pathways: the mPFC→BLA pathway preferentially encodes anxiety-related versus non-anxiety states, whereas the mPFC→ NAc pathway more selectively distinguishes exploratory from anxiety-like behaviors.

The above differentiation in population activity was not revealed by averaging the transient rate. When considering averaging the transient rate across a population can potentially mask characteristics unique to neurons that represent the emotional states of mice associated with these behaviors. We thus narrowed our focus to neurons that exhibited increased activity during heightened anxiety moments, such as entering the center zone. Termed “center-ON” neurons, these neurons displayed substantially increase in activity surrounding the mouse’s entry into the center zone, suggesting sensitivity to a behavioral shift from a low- to high-anxiety-like state triggered by entry into the anxiogenic center area (Fig. 2A). This increase in activity of the center-ON neurons in the mPFC→BLA pathway was not observed when the mice entered corner zones or engaged in sniffing or grooming (Fig. 2B). Upon comparing the percentage of center-ON neurons in the two pathways, we observed a higher prevalence of center-ON neurons in the mPFC→BLA group compared to the mPFC→NAc group (Fig. 2C and Supplementary Fig. 2). This finding supports the earlier observation that neurons in the mPFC→BLA pathway exhibit more pronounced encoding of anxiety-like states compared to those in the mPFC→NAc pathway. The distribution of center-ON neurons within the mPFC→BLA and mPFC→NAc pathways was consistent across animals and spanned both superficial and deep layers, as the 1-mm GRIN lens encompassed the full depth of the prelimbic cortex, indicating that their activity was not layer-specific. When we projected the activity patterns of these center-ON neurons onto the spatial distribution of the open field arena (Fig. 2D), both sets of center-ON neurons displayed heightened activity in the center zone. However, the center-ON neurons in the mPFC→NAc population exhibited increased activity during sniffing as well, which was not observed in the mPFC→BLA population (Fig. 2E and Supplementary Fig. 3). This finding implies that these two sets of center-ON neurons serve distinct roles in encoding behaviorally defined affective states in mice. Specifically, the mPFC→NAc neurons appear to be more involved in positively valenced behavioral states associated with exploration, while the mPFC→BLA neurons are more active during negatively valenced states linked to anxiety-like behavior. To reinforce this notion, we conducted the PCA exclusively on the center-ON neuron populations (Fig. 2F). Intriguingly, PCA scatter plots reveal behavioral clustering and spatial separations, demonstrating that different behaviors influence neuronal engagement differently in these two pathways. Again, we quantified the distances of neuronal activity patterns associated with Center, Sniffing, and Grooming from Corner behavior in both pathways (Fig. 2G). The result showed significant differences in the encoding of behaviorally defined affective states by center-ON ensembles in the two pathways. Similar to Fig. 1J, the greater separation between Center and Corner behaviors in the mPFC→BLA pathway suggests that this projection may play a critical role in modulating stress- and anxiety-related responses, underscoring its potential relevance as a therapeutic target for anxiety disorders. In contrast, the mPFC→NAc pathway appears more broadly engaged in general exploratory behaviors and less specialized in processing anxiety-specific signals (Fig. 2F and G). The presence of distinct center-ON neuronal populations further highlights the functional divergence between these pathways, with the mPFC→BLA and mPFC→NAc circuits exhibiting pathway-specific encoding of aversive versus exploratory behavioral states.

Distinct encoding of emotional status by center-ON neurons in the mPFC pathways

(A) Ca2+ traces of the averaged activity of center-ON neuron ensembles around the onset of center entry (5 s before to 5 s after). The total number of 24 center-ON neurons (n = 24) recorded in a mouse of the mPFC→BLA groups that used to calculate the representative averaged trace. Solid lines represent the averaged value. (B) Representative heatmaps depicting the activity patterns of mPFC→BLA neurons across four different behavioral contexts: Center, Corner, Sniffing, and Grooming. Each column corresponds to a different behavior, with the intensity of color indicating the level of neuronal activity. (C) Left: spatial distributions of center-ON neurons among the mPFC→BLA and mPFC→NAc neurons in one representative mouse from each group. Right: quantification of the percentage of center-ON neurons showed that the BLA pathway contained a significantly higher proportion of center-ON neurons compared to the NAc pathway. * P < 0.05, Mann-Whitney U test (Mann–Whitney U = 47, P = 0.049). (D) Spatial heatmaps of neuronal activity across the open field arena of all observed center-ON neurons in one example mouse of mPFC→BLA and mPFC→NAc groups. The color bar indicates the averaged normalized z-score. (E) Normalized transient rate of neuronal activity across behavioral states for mPFC→BLA and mPFC→NAc neurons. A two-way mixed-design ANOVA revealed a significant Region × Behavior interaction (P = 0.048), indicating that mPFC→NAc neurons exhibited higher activity during sniffing compared to mPFC→BLA neurons, while no regional differences were found in center, corner, or grooming states. * P < 0.05; two-way mixed-design ANOVA followed by Tukey’s post hoc test. (F) PCA plots illustrating the distribution of neuronal activity patterns during different behaviors for mPFC→BLA and mPFC→NAc pathways. Points are color-coded by behavior type (Center: blue, Corner: red, Sniffing: green, Grooming: purple). (G) Summary of the distances from the center of Corner behavior in terms of neuronal activity for each behavior in both mPFC pathways. ** P < 0.01, *** P < 0.001, Mann-Whitney U test.

Distinct activation of center-ON Neurons during anxiety-like behavioral states

To further validate our previous findings, we conducted the Elevated Plus Maze (EPM) test with the same mice one day following the OFT, as shown in Fig. 3A. The EPM test assesses anxiety by noting whether mice explore the exposed and elevated open arms (indicative of higher anxiety) or prefer the sheltered, enclosed arms (indicative of lower anxiety) (16). Our observations revealed that mice predominantly stayed in the closed arms, particularly favoring one closed arm (preferred closed arm, closed-P) over the other (non-preferred closed arm, closed-NP), suggesting a strong preference for a safer environment, despite both closed arms being designed to offer security (Fig. 3B). Neuronal activity in the mPFC→BLA and mPFC→NAc pathways was recorded during the EPM test. The neurons exhibited increased transient rates in the open arms, indicative of an anxiety-like behavioral response, though no significant differences were observed in the overall activity patterns between these pathways (Fig. 3C and Supplementary Fig. 4). This is consistent with earlier results from the OFT (Fig. 1H). Subsequent analyses focused on “center-ON” neurons previously identified during the OFT, examining their responses across different behavioral states in the EPM (Fig. 3D and E). Notably, center-ON neurons in the mPFC→BLA pathway demonstrated significantly heightened activity in the open arms compared to the closed arms, particularly when mice entered the open arms (Fig. 3F). The calcium transient rates for these neurons were markedly higher than those in the mPFC→NAc pathway, especially when mice were in the less familiar open-NP arm, suggesting that the unfamiliarity of this arm enhances anxiety-like behavior, and the mPFC→ BLA pathway is particularly sensitive to this affective shift (Fig. 3G and Supplementary Fig. 5). These findings corroborate our initial observations from the OFT, highlighting the critical role of mPFC→BLA neurons in modulating responses to anxiety-like behavioral states in mice. The pronounced activity differences in these neurons suggest their heightened sensitivity to anxiogenic stimuli and their potential involvement in perceiving and processing such states. Together, the results indicate that center-ON neurons in the mPFC→BLA pathway may play a more prominent role in detecting and potentially regulating anxiety-like behaviors compared to those in the mPFC→NAc pathway.

The mPFC→NAc and mPFC→BLA neurons demonstrated distinct activity patterns across various emotional states during EPM test.

(A) Schematic illustration of the Elevated Plus Maze (EPM) test setup showing the layout of open (Open-P, Open-NP) and closed (Close-P, Close-NP) arms used to assess anxiety-related behaviors in mice. (B) Box plots displaying the percentage of time spent by mice in the different arms of the EPM (Open-P, Open-NP, Close-P, Close-NP) for both mPFC→BLA and mPFC→NAc pathways. Data indicate variations in time spent across different arms, highlighting behavioral preferences. (C) Averaged transient rates (Hz) of neuronal activity of all recorded neurons in the mPFC→BLA and mPFC→NAc pathways in each arm of the EPM. Data are represented as mean ± SEM. Mixed-effects analysis (Region × Arm) showed higher transient rates in open arms relative to closed arms (Open-P vs Close-NP: β=0.0149, P=7.3×10⁻⁶; Open-NP vs Close-NP: β=0.0107, P=0.0019) with no overall Region effect (p=0.304) and no Region × Arm interaction (all P>0.29), indicating the activity pattern of the mPFC→BLA and mPFC→NAc pathways in each arm of the EPM is similar. (D) Analysis design of observing the neural activities of center-ON ensembles of OFT in different EPM arms. (E) Mean Ca2+ signal across the EPM arms of all observed center-ON neurons in one example mouse. The color bar indicates the averaged normalized z-score. (F) Example 30 center-ON neurons (top) and corresponding averaged activity (bottom) around the onset of arm entry during the EPM test. Time 0 represents the onset of arm entry. (G) Group data showing the mean transient rate of center-ON neurons across four EPM arms for mPFC→BLA and mPFC→NAc pathways. A two-way mixed-effects ANOVA with Region (BLA vs NAc) as a between-subject factor and Condition (Opent-P, Open-NP, Close-P, Close-NP) as a within-subject factor revealed a significant Region × Condition interaction (F(3,72) = 3.21, P = 0.030).Tukey’s post hoc tests showed that BLA exhibited significantly higher transient rates than NAc during the Open-NP condition (P = 0.018), while no significant regional differences were found in other three arms (P > 0.1). Data are represented as mean ± SEM. n = 9–13 mice for mPFC→BLA; n = 8-11 mice for mPFC→NAc.

The mPFC→NAc neurons, as opposed to mPFC→BLA neurons, encode emotional information associated with social preference

Social behavior can serve as an effective indicator of diverse affective states in individuals (17) and is encoded by the mPFC neuron ensembles with distinct activity patterns (18). We thus investigated whether distinct activity patterns are exhibited by mPFC→BLA and mPFC→NAc neurons, correlating with different socially relevant behavioral states. The social test was conducted using a three-chamber apparatus (Fig. 4A). The two end chambers were equipped with either another mouse or an object. The subject mouse underwent a 10-minute acclimation phase in the central 45 cm chamber, with the 10 cm end compartments empty. Subsequently, three ten-minute testing sessions (S1, S2, and S3) were conducted, each involving different stimuli in the end chambers to assess interaction. Utilizing automated tracking software, the time spent within the “sniffing zone” was measured. Notably, the mice exhibited a preference for interacting with fellow mice over objects in S1 and S2. In S2, this preference gap narrowed due to the relocation of objects. During S3, when the choice was between two mice, the test subjects demonstrated a significant preference for exploring the novel mouse (Fig. 4B), consistent with established patterns of social novelty behavior. These findings collectively signify a normal level of sociability as observed in the three-chamber test.

Pattern decorrelation is shown by the mPFC→NAc neurons, but not mPFC→BLA neurons.

(A). Top: the apparatus for the social interaction test. The subject mouse is in the middle 45 cm-long chamber; the 10 cm end compartments contain different stimuli. Bottom: the three sessions of the social interaction test. In the first 10-minute test session (S1), a strange mouse (M1) and an object (O) were placed in the end chambers; session 2 (S2) used the same stimuli but swapped their positions; in session 3 (S3) a new mouse (M2) replaced O, so that the subject mouse must choose whether to interact with a familiar or strange mouse (M1 vs. M2). (B). The percentage of time that mice (n = 10) spent interacting with stimuli in each session. Data are shown as mean ± SEM. Paired t-test, S1: t(9) = 12.71, P = 1.2×10⁻⁶, Cohen’s d = 4.02; S2: t(9) = 0.88, P = 0.40, Cohen’s d = 0.28; S3: t(9) = 7.29, P = 6.7×10⁻5, Cohen’s d = 2.31. Mice spent significantly more time with the social mouse (M1) than with the object (O) in S1 or the new social mouse (M2) than with the old social mouse (M1) in S1. (C). The averaged calcium event rate of mPFC→BLA group when mice were engaged with different stimuli over the three sessions. n = 13 mice. Data are represented as mean ± SEM. Paired two tailed t-test, S1: t(12) = -0.77, P = 0.46, Cohen’s d = -0.22; S2: t(12) = 033., P = 0.74, Cohen’s d = 0.09; S3: t(12) = -0.39, P = 0.070, Cohen’s d = - 0.11. The transient rate of the mPFC→BLA neurons showed no significant difference during the interactions with the two stimuli in each session. (D). The averaged calcium event rate of mPFC→NAc group when mice were engaged with different stimuli over the three sessions. n = 11 mice. Data are represented as mean ± SEM. Paired two tailed t-test, S1: t(10) = -0.88, P = 0.40, Cohen’s d = -0.27; S2: t(10) = 178, P = 0.11, Cohen’s d = 0.56; S3: t(10) = -1.45, P = 0.18, Cohen’s d = 0.46. The transient rate of the mPFC→NAc neurons showed no significant difference during the interactions with the two stimuli in each session. (E-F). Top: raster plots of correlation coefficient of paired mPFC neurons during interactions with different stimuli, from a representative mouse in mPFC→BLA and mPFC→NAc group, respectively. Bottom: distribution of pair-wise Pearson correlation coefficients among these recorded mPFC neurons in responding to stimuli in each session. The blue and orange plots represent M1 interactions, while the lighter colors represent the interaction with the other stimulus (O). (G). The averaged full width at half maximum (FWHM) of correlation coefficient distribution of individual mice in mPFC→BLA (n = 8) group during interactions with different stimuli Data are shown as mean ± SEM. Paired two tailed t-test, S1: t(7) = -1.63, P = 0.15, Cohen’s d = -0.58; S2: t(7) = -0.41, P = 0.69, Cohen’s d = -0.14; S3: t(7) = -1.63, P = 0. 15, Cohen’s d = 0.57. The correlation coefficient of the mPFC→BLA neurons showed no significant difference between two stimuli. (H). The averaged full width at half maximum (FWHM) of correlation coefficient distribution of individual mice in mPFC→NAc (n = 8) group during interactions with different stimuli Data are shown as mean ± SEM. *P<0.05, Paired two tailed t-test, S1: t(7) = -2.75, P = 0.025, Cohen’s d = -0.92; S2: t(7) = -0.92, P = 0.39, Cohen’s d = - 0.33; S3: t(7) = -2.87, P = 0.021, Cohen’s d = -0.96. The correlation coefficient of the mPFC→NAc neurons showed significantly stronger separation between the two stimuli during S1 and S3 than during S2.

To explore the neuronal activity associated with distinct social states, we closely monitored the excitatory mPFC→BLA and mPFC→NAc neurons within the prelimbic region of the mPFC. Surprisingly, both the mPFC→BLA and mPFC→NAc neurons exhibited little discernible difference in their Ca2+ transient rate while the mice interacted with the social stimulus (mouse 1, M1) and the nonsocial stimulus (object, O) (Fig. 4C and D). Likewise, no significant variation was observed during S3. Given that interactions with social versus nonsocial stimuli are known to elicit distinct affective states, as reflected in social preference behavior, our findings suggest that the neural encoding of these states is not solely dependent on the overall firing rate of mPFC→BLA and mPFC→NAc neurons.

Next, we analyzed the ability of these two groups of neurons to convey information through coactivity patterns (19). Functional correlations (coactivity) within the circuit can be identified using Pearson correlation coefficients, and pattern decorrelation serves to clarify overlapping activity patterns. This mechanism is akin to how the olfactory bulb distinguishes closely related odorants based on minor structural variations (20). In a previous study, we demonstrated that pattern decorrelation of mPFC excitatory circuit neuronal activities is indispensable for social preference (21). This led us to contemplate whether differences existed between the mPFC→BLA and mPFC→NAc neurons. We thus calculated the pairwise Pearson correlation coefficients of neuronal activities specifically during the interaction with the social or object stimulus (Fig. 4E and F). The pairwise correlation coefficients for different stimuli across the recorded mPFC→BLA and mPFC→NAc neurons from all mice exhibited a distribution around zero, albeit with varying widths. Intriguingly, we observed a significant pattern decorrelation only within the mPFC→NAc neurons, evident by the full-width at half maximum (FWHM) of the correlation distribution being lower for the more-attractive stimuli in each session (Fig. 4F and H). In contrast, such decorrelation was not observed in the mPFC→BLA neurons (Fig. 4E and G). The broader tuning profile of mPFC→NAc neurons suggests a more sustained and robust encoding of social preference, reflected in distinct coactivity patterns across ensembles. While this enhanced encoding aligns with the positive valence often associated with social interactions, it does not by itself establish direct encoding of emotional states. Instead, these findings indicate that the mPFC→NAc pathway may play a prominent role in processing and guiding social behavior, whereas mPFC→BLA neurons appear less engaged in differentiating between social and nonsocial stimuli at the level of ensemble coactivity.

We next explored the distinct activity patterns of mPFC→BLA and mPFC→NAc neurons during social (M1) and nonsocial (object, O) interactions by comparing all recorded neurons and the subset of center-ON neurons previously identified (Fig. 5). The PCA analysis for all recorded neurons in both pathways showed that the mPFC→BLA neurons exhibit greater separation between social (M1) and nonsocial (O) stimuli than the mPFC→NAc neurons (Fig. 5A). This indicates that mPFC→BLA neurons display a more distinct coding of social and nonsocial interactions. We quantified the distance between the neuronal activity clusters for social and nonsocial stimuli, showing a significantly greater separation in the mPFC→BLA pathway compared to the mPFC→NAc pathway, supporting the notion that mPFC→BLA neurons are more attuned to distinguishing between these stimuli (Fig. 5B). We then focused specifically on center-ON neurons identified during the OFT. The PCA plots showed that the center-ON neurons in the mPFC→NAc pathway exhibited more distinct separation between social and nonsocial stimuli compared to the mPFC→BLA neurons which showed much less differentiation (Fig. 5C). We further quantified the distances between social and nonsocial stimuli for center-ON neurons, with mPFC→NAc neurons showing significantly greater separation compared to mPFC→BLA neurons (Fig. 5D). Overall, the results demonstrate that while mPFC→BLA neurons exhibit broader population-level encoding of social versus nonsocial stimuli, the mPFC→NAc pathway is more specialized in differentiating these stimuli within behaviorally relevant neuronal subsets, such as center-ON neurons. This distinction highlights complementary roles for these pathways: mPFC→BLA neurons contribute to general contextual or environmental differentiation, whereas mPFC→NAc neurons are more specifically involved in encoding social preference.

Comparative analysis of all neurons and center-ON subsets reflects divergent encoding patterns in the mPFC→BLA and mPFC→NAc Pathways.

(A) PCA plots showing the activity of all recorded neurons in the mPFC→BLA and mPFC→NAc pathways during interactions with a social stimulus (mouse, M1) and a nonsocial stimulus (object, O) in a modified Three Chamber Test (mTC). The plots reveal distinct clustering of neural activity patterns for each stimulus within both pathways. (B) Distances between all neuronal activity clusters for nonsocial (O) interactions from the center of that for M1 interaction, compared across mPFC→BLA and mPFC→NAc pathways. Significant difference is demonstrated in the mPFC→BLA pathway. (C) PCA plots for center-ON neurons identified during the Open Field Test (OFT). These plots illustrate the activity of these neurons in the mPFC→BLA and mPFC→NAc pathways during the same social and nonsocial stimuli. Clusters show how Center-ON neurons specifically respond to each type of stimulus. (D) Distances of all center-ON neuronal activity clusters from the center of that for M1 interaction, compared across mPFC→BLA and mPFC→NAc pathways. Significant difference is demonstrated in the mPFC→NAc pathway. Data are represented as mean ± SEM. ***P < 0.001, Mann-Whitney U test.

To evaluate the consistency of our neuron co-registration across behavioral tests, we quantified the proportion of neurons reliably matched between sessions using the open-field test (OFT) as the reference. For the mPFC→BLA pathway, the mean yield of matched neurons was approximately 30% between OFT and EPM (range: 5-92%, n = 19 mice) and 16% between OFT and Social tests (range: 5-32%, n = 14 mice). For the mPFC→NAc pathway, the mean yields were 23% (range: 11-71%, n = 20 mice) and 21% (range: 13-37%, n = 17 mice) for OFT-to-EPM and OFT-to-Social alignments, respectively. Despite variability across animals, these values indicate that a substantial subset of neurons can be reliably tracked across days and behavioral contexts, supporting the stability of our longitudinal imaging alignment method.

The mPFC→NAc and the mPFC→BLA neurons were differentially modulated by chronically induced affective states

Demonstrating that mPFC→NAc and mPFC→BLA neurons are preferentially engaged during acute positively valenced (exploratory) and negatively valenced (anxiety-like) behavioral states, respectively, we next asked whether these pathways are also modulated by naturally induced affective states. Social competition outcomes, such as winning or losing, are ecologically valid means of eliciting opposing emotional experiences in rodents. The social dominance tube test has been widely used to assess social hierarchy and aggression (22, 23), a critical component of social stress and emotion regulation in rodents. It may therefore serve as an effective model for natural emotion induction. Studies have shown that social defeat or subordination can lead to measurable changes in neural activity, stress hormone levels, and behavior (24–27). Therefore, we believe that while the tube test may not induce long-lasting emotional changes on its own, it can serve as an effective and ethologically relevant model for natural emotion induction, particularly in the context of transient emotional responses associated with social competition and status. In this assay, two mice are placed at opposite ends of a narrow transparent tube and allowed to interact as they attempt to pass through to the other side. (22, 23, 28, 29). Repeated daily over a week, this procedure often results in mice adopting consistent winner or loser roles, potentially leading to sustained positively or negatively valenced affective states (Fig. 6A). To control for potential time-dependent or nonspecific effects, a group of singly housed mice that did not undergo tube testing but were handled under identical conditions and on the same experimental timeline served as controls. Corticosterone levels measured from brain lysates revealed that loser mice exhibited significantly higher concentrations than both control and winner groups, indicating that the observed physiological changes were specific to social subordination rather than general effects of time or repeated exposure (Fig. 6B).

The modifications in social ranking of mice alter their anxiety and social states.

(A) Schematic of the experimental timeline. Singly housed mice were handled for 3 days prior to undergoing a 5-day tube test to establish dominant (Winner) and subordinate (Loser) social status. Naïve singly housed mice served as controls. Behavioral assays were conducted 24 hours after, and brain tissue collection within 1 hour of, the final tube test session. (B) Corticosterone concentrations measured from brain lysates of Control, Winner, and Loser mice. Loser mice showed significantly elevated corticosterone levels compared to both Control and Winner groups (** P < 0.01, # P < 0.05). Data are presented as mean ± SEM. (C-D) The mice’s total distance traveled in the open field and time spent in the center zones before and after the tube test. *P < 0.05; ns, non-significant. n = 6 mice for each group, winner vs loser, unpaired Student’s t-test. (E) The differences of total distance (top) and center time (below) between winner and loser groups before and after the tube tests. **P < 0.01; ns, non-significant. Before vs. after tube test, Two-way ANOVA Multiple comparisons, sidak’s post-hoc test. (F) The mice’s exploration time near to M1 and O/M2 during the social ability test (left) and social memory test (right) conducted before the tube test. Social ability index calculated as (time near M1 chamber - time near object chamber) / (time near M1 chamber + time near object chamber). Social memory index calculated as (time near M2 chamber - time near M1 chamber) / (time near M2 chamber + time near M1 chamber). *P < 0.05, ***P < 0.001, ****P < 0.0001, ns, non-significant. Two-way ANOVA Multiple comparisons, sidak’s post-hoc test. Unpaired Student’s t-test for preference index. (G) The mice’s exploration time spent near the M1 and O/M2 during social ability test and social memory test after the tube test. *P < 0.05, **P < 0.01, ***P < 0.001, ****P < 0.0001. Two-way ANOVA Multiple comparisons, sidak’s post-hoc test. Unpaired Student’s t-test for preference index. (H) The comparisons of the social ability index (top) and social memory index (below) of winner and loser group before and after the tube tests. *P < 0.05, **P < 0.01, ns, non-significant. Two-way ANOVA Multiple comparisons, sidak’s post-hoc test). n = 6 mice for each group.

To verify if the tube test genuinely altered emotional states, littermate mice underwent an open field test both before and after the tube test (Fig. 6C-E). Initially, there was no difference in the travel distance or time spent in the center of the open field arena between winner and loser mice (Fig. 6C). Post-tube test, travel distances remained unchanged, suggesting that mobility was not affected by induced emotions (Fig. 6D and E). However, loser mice spent significantly less time in the center, indicating heightened anxiety and successful negative emotion induction. In contrast, winners’ time in the center remained unchanged, implying their anxiety levels remained normal (Fig. 6D and E). Furthermore, to determine if the tube test could evoke positive emotions in winner mice, their sociability was assessed before and after the test because positive emotion (such as the pleasure from winning) would enhance sociability. Initially, both the winner and loser mice showed a standard preference for interacting with a live mouse over an inanimate object, and a new mouse over a familiar one (Fig. 6F). After the tube test, winners displayed a higher sociability index towards social mice and a tendency for greater interest in new mice. In contrast, losers almost entirely lost their social preference, underscoring the induction of negative emotions (Fig. 6G and H). These findings confirm that repeated tube test competitions effectively induced distinct affective states, with winners exhibiting positively valenced responses and losers displaying negatively valenced responses.

Earlier version: Next, we compared the synaptic strength in the mPFC→NAc and mPFC→BLA circuits following emotion induction via the repeated tube test. To activate these circuits, we used AAV-mediated expression of channelrhodopsin-2 (ChR2) for selective stimulation of prelimbic (PL) neurons in the NAc and BLA using blue light. ChR2-EYFP expression was verified in the PL, with terminal expression in the NAc and BLA (Fig. 7A and B). Following the establishment of stable social rankings through the tube test (Fig. 7C), we recorded from brain slices containing the NAc or BLA. Blue light application confirmed ChR2 activation, triggering inward currents and voltage spikes in both winner and loser mice (Fig. 7D and H). In the mPFC→NAc pathway, blue light stimulation evoked larger excitatory postsynaptic currents (EPSCs) in winner mice compared to losers (Fig. 7E), despite similar paired-pulse ratios (PPRs) in both groups (Fig. 7F and G). This suggests stronger synaptic transmission in winners’ mPFC→NAc circuits. Conversely, in the mPFC→BLA pathway, loser mice exhibited larger EPSCs and significantly higher PPRs than winners (Fig. 7I-K), indicating stronger synaptic transmission in losers’ mPFC→BLA circuits. Indeed, we found that the AMPA/NMDA receptor-mediated EPSC ratio was higher in winners’ mPFC→NAc circuits and losers’ mPFC→BLA circuits (Fig. 7L and M). This result showed that winning and losing influence distinct neural pathways, with winners enhancing synaptic transmission in the mPFC→NAc circuit and losers in the mPFC→BLA circuit.

Revised version: Next, we examined activity-dependent differences in light-evoked synaptic responses within the mPFC→NAc and mPFC→BLA circuits following emotion induction via the repeated tube test. To selectively activate these pathways, we used AAV-mediated expression of channelrhodopsin-2 (ChR2) in prelimbic (PL) neurons and stimulated their terminals in the NAc or BLA with blue light. ChR2-EYFP expression was verified in the PL, with terminal expression in the NAc and BLA (Fig. 7A and B). Following the establishment of stable social rankings through the tube test (Fig. 7C), we recorded from brain slices containing the NAc or BLA. Blue light stimulation reliably activated ChR2-expressing inputs, eliciting inward currents in both winner and loser mice (Fig. 7D and H). We note that differences in light-evoked response amplitudes across slices may reflect variability in opsin expression or effective light delivery. To mitigate these factors, responses were normalized within each slice, and LED pulse duration was adjusted as needed to reliably evoke activity. Thus, our observations reflect relative circuit responsiveness under standardized recording conditions rather than absolute synaptic strength across animals. In the mPFC→NAc pathway, light-evoked excitatory postsynaptic currents (EPSCs), analyzed under normalized stimulation conditions, showed relatively more sensitive responses in winner mice compared to losers (Fig. 7E). These results indicate an enhanced functional sensitivity of the mPFC→NAc circuit to optogenetic stimulation in winner mice. Moreover, we examined EPSC kinetics in postsynaptic neurons from winner and loser mice. Both the mEPSC rise and decay slopes were comparable between the two groups (Fig. 7F). In addition, paired-pulse ratios (PPRs) did not differ between groups (Fig. 7G and H). Conversely, in the mPFC→BLA pathway, loser mice exhibited relatively larger normalized EPSCs (Fig. 7J). Interestingly, both the EPSC rise and decay slopes were significantly increased in loser mice compared with winners, indicating enhanced synaptic responsiveness and accelerated charge transfer kinetics in BLA neurons of loser (Fig. 7K). Moreover, PPRs were significantly higher in loser mice than in winners (Fig. 7L and M), suggesting increased presynaptic facilitation and heightened functional sensitivity of this pathway in losers. Consistent with these observations, the AMPA/NMDA EPSC ratio was relatively higher in the mPFC→NAc pathway of winners and in the mPFC→BLA pathway of losers (Fig. 7N and O). Together, these findings indicate that winning and losing differentially bias the functional engagement of distinct prefrontal output pathways under optogenetic activation, with winners preferentially engaging the mPFC→Nac circuit and losers the mPFC→BLA circuit.

The social status dependent PL-NAc and PL-BLA neuronal activities.

Earlier version: (A) Schematic showing the AAV-ChR2 injection in the PL of mPFC. (B) Expression of ChR2-EYFP in the mPFC, and terminal expression of ChR2-EYFP in the NAc and BLA. Scale bar, 400 μm, 50 μm, and 50 μm, from left to right. (C) The ChR2-EYFP expressed mice underwent tube test to determine their winner and loser status. (D) The EPSCs in the NAc were evoked by 470 nm blue light. Upper, schematic showing the light fiber placement and EPSC recording sites in the NAc; Lower, representative EPSCs in the brain slices from winner (red) and loser (black) mice. Scale bar, 30 pA and 25 ms. (E) The normalized light intensity-dependent EPSCs. The data points were normalized to 10 μW and were shown as mean value. Winner, n= 19 neurons/ 7 mice; Loser n= 14 neurons/ 7 mice. (F) The representative pair pulse ratio (PPR) of the EPSCs in the NAc (the light intensity were 1-2 μW and the interpulse interval were 50 and 75 ms). Scale bar, 8 pA and 25 ms. (G) Summarized plots of the PPR under different interpulse intervals. Winner, n= 7 neurons/ 7 mice; Loser, n = 14 neurons/ 7 mice. ns, non-significant. Scale bar, 8 pA and 25 ms. (H-I) EPSC recordings in the BLA. Winner, n = 17 neurons/ 7 mice; Loser n = 21 neurons/ 7 mice. Scale bar, 6 pA and 25 ms. (J) The representative PPR of the EPSCs in the BLA. (K) Summarized plots of the PPR under different interpulse intervals. Winner, n = 8 neurons/ 7 mice; Loser, n = 11 neurons/ 7 mice. Scale bar, 3 pA and 25 ms. (L) Left, representative EPSCs in the PL-NAc neurons from winner (red) and loser (black) mice. Right, the summarized AMAP/NMDA ratio. Winner, n= 5 neurons/ 5 mice; Loser n= 7 neurons/ 7 mice, 1-5 ms pulse). Scale bar, 25 pA and 50 ms. (M) Left, representative EPSCs in the PL-BLA neurons from winner (red) and loser (black) mice. Right, the summarized AMAP/NMDA ratio. Winner, n= 12 neurons/ 7 mice; Loser n= 21 neurons/ 7 mice, 1-10 ms pulse). Scale bar, 10 pA and 50 ms. Data are shown as mean ± SEM. *P < 0.05, **P < 0.01, ***P < 0.001. Two-way ANOVA for E, G, I, and K; two-tailed paired Student’s t-test for L and M.

Revised version: (A) Schematic showing the AAV-ChR2 injection in the PL of mPFC. (B) Expression of ChR2-EYFP in the mPFC, and terminal expression of ChR2-EYFP in the NAc and BLA. Scale bar, 400 μm, 50 μm, and 50 μm, from left to right. (C) The ChR2-EYFP expressed mice underwent tube test to determine their winner and loser status. (D) The EPSCs in the NAc were evoked by 470 nm blue light. Upper, schematic showing the light fiber placement and EPSC recording sites in the NAc; Lower, representative EPSCs in the brain slices from winner (red) and loser (black) mice. Scale bar, 30 pA and 25 ms. (E) The normalized light intensity-dependent EPSCs. The data points were normalized to 10 μW and were shown as mean value. Winner, n = 19 neurons/ 7 mice; Loser n = 14 neurons/ 7 mice. (F) Analysis of EPSC rise and decay slopes in the NAc. Data are shown as median + Interquartile Range. Rise slope: ns, non-significant, P = 0.0780, n= 20 neurons/ 7 mice; Decay slope: ns, P = 0.0732, n= 17 neurons/ 7 mice. Non-parametric testing were analyzed using the Mann-Whitney test, and effect size was reported as Rank-biserial correlation (r). (G) The representative pair pulse ratio (PPR) of the EPSCs in the NAc (the light intensity were 1-2 μW and the interpulse interval were 50 and 75 ms). Scale bar, 8 pA and 25 ms. (H) Summarized plots of the PPR under different interpulse intervals. Winner, n= 7 neurons/ 7 mice; Loser, n = 14 neurons/ 7 mice. ns, non-significant. Scale bar, 8 pA and 25 ms. (I-J) EPSC recordings in the BLA. Winner, n = 17 neurons/ 7 mice; Loser n = 21 neurons/ 7 mice. Scale bar, 6 pA and 25 ms. (K) Analysis of EPSC rise and decay slopes in the BLA. Data are shown as median + Interquartile Range. Rise slope: *P = 0.0138, n=15; Decay slope: *P = 0.0392, n = 19. The difference remained significant using non-parametric testing (Mann-Whitney U test), and effect size was reported as Rank-biserial correlation (r). (L) The representative PPR of the EPSCs in the BLA. (M) Summarized plots of the PPR under different interpulse intervals. Winner, n = 8 neurons/ 7 mice; Loser, n = 11 neurons/ 7 mice. Scale bar, 3 pA and 25 ms. (N) Left, representative EPSCs in the PL-NAc neurons from winner (red) and loser (black) mice. Right, the summarized AMAP/NMDA ratio. Winner, n= 5 neurons/ 5 mice; Loser n= 7 neurons/ 7 mice, 1-5 ms pulse). Scale bar, 25 pA and 50 ms. (O) Left, representative EPSCs in the PL-BLA neurons from winner (red) and loser (black) mice. Right, the summarized AMAP/NMDA ratio. Winner, n= 12 neurons/ 7 mice; Loser n= 21 neurons/ 7 mice, 1-10 ms pulse). Scale bar, 10 pA and 50 ms. Data are shown as mean ± SEM except for panels (F) and (K). *P < 0.05, **P < 0.01, ***P < 0.001. Two-way ANOVA for E, G, I, and K; two-tailed paired Student’s t-test for L and M.

Removed in the revised version: Importantly, we observed no significant differences in synaptic strength between winner and loser mice within the mPFC region. The correlation between the light intensity and the normalized amplitude of blue light-induced current showed no variance between the winner and loser groups during light stimulation and recordings in the PL region of the mPFC (Supplementary Fig. 6A-C). Furthermore, the induction of emotional states did not alter the excitability of mPFC neurons; this is evident from the consistent relationship between light stimulation frequency and the number of spikes induced across mPFC neurons from both winner and loser mice (Supplementary Fig. 6D and E). Consequently, it appears that the induction of emotion specifically affects the efferent pathway from the mPFC, rather than the neuronal activity within the mPFC itself. Overall, our findings reveal that chronically induced positively valenced affective states are associated with enhanced synaptic strength in the mPFC→NAc pathway, whereas negatively valenced states correlate with increased synaptic strength in the mPFC→BLA pathway. These results underscore the divergent modulation of mPFC projections to the NAc and BLA by affective states of opposing valence, further highlighting the distinct roles of these circuits in encoding emotionally salient experiences.

Discussion

This study investigates the distinct roles of mPFC projections to the BLA and NAc in processing affective and social stimuli. Using a combination of behavioral paradigms and in vivo neural recordings, we found that although both pathways contribute to the encoding of emotionally salient and socially relevant behaviors, they exhibit divergent activity dynamics and functional specializations. We interpret the observed behaviors across the OFT, EPM, and tube test as validated proxies for internal emotional states, based on extensive prior literature. For example, avoidance of the center area in the open field test and the open arms in the elevated plus maze are widely accepted as behavioral indices of anxiety-like states in rodents (14, 16). Similarly, outcomes of the tube test have been shown to predict social hierarchy and are associated with long-lasting emotional valence shifts, including increased stress vulnerability in subordinate (loser) animals (22, 23). Prior research has shown that transitions between exploratory and defensive behaviors are associated with rapid shifts in neural ensemble activity in the BLA, consistent with changes in internal affective or motivational states (30, 31). These studies support the use of well-defined behavioral states as reliable proxies for emotional dynamics. While direct measurement of physiological stress markers such as corticosterone would further strengthen this link, such measurements are not feasible during in vivo calcium imaging, particularly during rapid state transitions in freely moving mice. Therefore, our interpretation of circuit activity in relation to “emotional states” relies on a validated behavioral framework supported by ensemble-level neural correlates.

Earlier version: The mPFC→BLA pathway demonstrated a broader capacity to differentiate between a range of affective states, particularly those associated with anxiety-like behavior, suggesting a generalized role in emotional processing of aversive or negatively valenced stimuli. In contrast, the mPFC→NAc pathway, especially its center-ON neurons, was more selectively tuned to exploratory behaviors and social interactions, indicating a role in encoding reward-related information and approach-oriented behavioral states (6). This differentiation is most evident in the context of anxiety-like behavior, where the mPFC→BLA neurons show enhanced responsiveness, highlighting their pivotal role in regulating anxiety-related states. This observation concurs with prior research highlighting the Amygdala’s pivotal role in the processing of fear and anxiety (32, 33).

Revised version: The mPFC→BLA pathway demonstrated a broader capacity to differentiate between a range of affective states, particularly those associated with anxiety-like behavior, suggesting a generalized role in emotional processing of aversive or negatively valenced stimuli. In contrast, the mPFC→NAc pathway, especially its center-ON neurons, was more selectively tuned to exploratory behaviors and social interactions, indicating a role in encoding reward-related information and approach-oriented behavioral states (6). This differentiation is most evident in the context of anxiety-like behavior, where the mPFC→BLA neurons show enhanced responsiveness, highlighting their pivotal role in regulating anxiety-related states. This observation concurs with prior research highlighting the pivotal role of the Amygdala in the processing of fear and anxiety (32, 33).

The contrasting patterns of neuronal activity observed between “center-ON” neurons and the entire neuronal population further underscore the functional diversity within the mPFC→BLA and mPFC→NAc pathways (Fig. 1 and 2). The broader population of neurons does not exhibit the same selective sensitivity to affective state changes as the “center-ON” neurons. Instead, their activity appears more generalized, lacking the specificity observed in the center-ON subset. This suggests that while the overall activity within these pathways may contribute to the modulation of emotional responses, it is the center-ON neurons that provide a more precise and dynamic representation of affective shifts, particularly those related to anxiety-like behavior. Acting as a specialized subpopulation, center-ON neurons appear finely attuned to salient triggers such as entry into anxiogenic environments, offering a more detailed encoding of emotionally relevant contexts. In contrast, the broader neuronal population may support general regulation of affective states, albeit in a less differentiated manner. Together, these findings highlight the complexity of emotional processing within mPFC circuits, revealing the coexistence of specialized and generalized encoding mechanisms that collectively shape nuanced behavioral responses to emotionally salient stimuli.

Based on our findings in the social behavior paradigm, a more nuanced picture emerges of how different neuronal subsets within the mPFC pathways respond to social and nonsocial stimuli. At the population level, the mPFC→BLA pathway shows significantly greater differentiation between social and nonsocial stimuli compared to the mPFC→NAc pathway (Fig. 5A, B). This suggests that mPFC→BLA neurons are broadly tuned to encode differences across diverse stimulus categories, supporting generalized responses to varied environmental contexts. Such a role may be particularly relevant because anxiety-like states can be triggered by many types of cues, whether social or nonsocial, requiring a more holistic mode of neural processing. In contrast, within the center-ON neuronal subset, the mPFC→NAc pathway exhibits significantly greater differentiation between social and nonsocial stimuli compared to the mPFC→BLA pathway (Fig. 5C, D). This finding suggests that this subset of mPFC→NAc neurons is more specialized in encoding motivationally salient social cues, potentially playing a critical role in shaping specific aspects of social behavior. Importantly, while social preference is often associated with positive valence, the sharper tuning observed in mPFC→NAc neurons should be interpreted as enhanced encoding of social preference itself, with implications for, but not a direct measure of, affective valence. Together, these results highlight a layered encoding scheme in which the mPFC→BLA pathway provides a broad, generalized framework for environmental differentiation, whereas the center-ON neurons in the mPFC→NAc pathway selectively refine distinctions between social and nonsocial stimuli. This dual strategy may enable both flexible adaptation to general environmental contexts and precise behavioral responses in socially nuanced situations.

Earlier version: Furthermore, our study reveals a compelling dimension of how chronic affective states influence synaptic strength within specific mPFC pathways. Sustained positively valenced states were associated with enhanced synaptic strength in the mPFC→NAc pathway (Fig. 7E), consistent with the NAc’s established role in reward processing and positive reinforcement (34, 35). In contrast, chronic negatively valenced states were linked to increased synaptic strength in the mPFC→BLA pathway (Fig. 7I-K), reinforcing the BLA’s critical role in processing aversive stimuli and mediating fear responses (36, 37). These findings underscore the influence of social status on the functional plasticity of mPFC efferent pathways involved in reward and emotional processing. This insight provides an important step toward understanding how hierarchical social dynamics shape emotional state representations at the circuit level. Nonetheless, further investigation is needed to uncover the underlying molecular and cellular mechanisms and to determine whether similar status-dependent effects are observed in other brain regions involved in affective regulation.

Revised version: Furthermore, our study reveals a compelling dimension of how chronic affective states influence synaptic strength within specific mPFC pathways. Sustained positively valenced states were associated with enhanced synaptic strength in the mPFC→NAc pathway (Fig. 7E), consistent with the NAc’s established role in reward processing and positive reinforcement (34, 35). In contrast, chronic negatively valenced states were linked to increased synaptic strength in the mPFC→BLA pathway (Fig. 7I-K), reinforcing the BLA’s critical role in processing aversive stimuli and mediating fear responses (36, 37). These findings underscore the influence of social status on the functional plasticity of mPFC efferent pathways involved in reward and emotional processing. This insight provides an important step toward understanding how hierarchical social dynamics shape emotional state representations at the circuit level. Nonetheless, further investigation is needed to uncover the underlying molecular and cellular mechanisms and to determine whether similar status-dependent effects are observed in other brain regions involved in affective regulation. While the tube test is traditionally used to determine social dominance rather than as a classical stress paradigm, repeated social competition produced measurable physiological changes, including increased corticosterone levels in loser mice, suggesting that the paradigm captures aspects of social hierarchy-related stress.

Earlier version: The current study focused on decoding the hidden variabilities in mPFC efferent pathways across emotional states of adult female mice, given that females are nearly twice as likely as males to be diagnosed with anxiety and depressive disorders (38, 39). Previous work (40) demonstrated that female rodents exhibit a higher density of dendritic spines on infralimbic mPFC neurons projecting to the BLA following fear conditioning, and show more rapid extinction of conditioned fear responses, suggesting structural and functional adaptations that differ from males. Stress responses differ between genders, with men showing increased activity in the prefrontal cortex and women in the limbic system (41). In female rodents, the neuronal firing activity of BLA shifts across the estrous cycle, which contributes to the sex-differences of BLA activity changes effected by repeated restraint stress treatment (42). Meanwhile, when receiving chronic variable stress, the functional and morphological changes of NAc-projecting mPFC neurons were more severe in female than in male mice (43). Finally, sex-specific transcriptional signatures occur in the PFC neurons of human patients with major depressive disorder (MDD), with females showing downregulation of DUSP6 gene associated with increasing pERK in pyramidal neurons, which impacts on excitatory synaptic transmission in PFC of depressed female or stressed mice. On the other hand, upregulation of EMX1 gene in PFC of male with MDD may increase their neuronal excitability and stress susceptibility (44). Taken together, these studies highlight pronounced sex differences in the structure, plasticity, and functional output of mPFC–BLA and mPFC–NAc circuits, which likely underlie the female-predominant prevalence of anxiety and depressive disorders. Future studies will investigate how mPFC efferent pathways dynamically code emotional states in male animals, thereby elucidating the circuit-level mechanisms of sexually dimorphic behaviors.

Revised version: The current study focused on decoding the hidden variabilities in mPFC efferent pathways across emotional states of adult female mice, given that females are nearly twice as likely as males to be diagnosed with anxiety and depressive disorders (38, 39). Previous work (40) demonstrated that female rodents exhibit a higher density of dendritic spines on infralimbic mPFC neurons projecting to the BLA following fear conditioning, and show more rapid extinction of conditioned fear responses, suggesting structural and functional adaptations that differ from males. Stress responses differ between genders, with men showing increased activity in the prefrontal cortex and women in the limbic system (41). In female rodents, the neuronal firing activity of BLA shifts across the estrous cycle, which contributes to the sex-differences of BLA activity changes effected by repeated restraint stress treatment (42). Meanwhile, when receiving chronic variable stress, the functional and morphological changes of NAc-projecting mPFC neurons were more severe in female than in male mice (43). Finally, sex-specific transcriptional signatures occur in the PFC neurons of human patients with major depressive disorder (MDD), with females showing downregulation of DUSP6 gene associated with increasing pERK in pyramidal neurons, which impacts on excitatory synaptic transmission in PFC of depressed female or stressed mice. On the other hand, upregulation of EMX1 gene in PFC of male with MDD may increase their neuronal excitability and stress susceptibility (44). Taken together, these studies highlight pronounced sex differences in the structure, plasticity, and functional output of mPFC-BLA and mPFC-NAc circuits, which likely underlie the female-predominant prevalence of anxiety and depressive disorders. One limitation of the present study is that the experiments were performed primarily in female mice. Sex differences in prefrontal-limbic circuits have been reported in several behavioral paradigms, and future studies will be needed to determine whether similar projection-specific coding principles extend to male animals. Investigating how mPFC efferent pathways dynamically code emotional states in male animals will help elucidate the circuit-level mechanisms of sexually dimorphic behaviors.

It is also important to recognize that the diversity of activity observed here may, in part, reflect underlying anatomical and cellular heterogeneity. The mPFC→NAc pathway comprises neurons from multiple cell classes, and mPFC→BLA projections traverse different cortical layers. Likewise, both the BLA (magnocellular vs. parvocellular subdivisions) and the NAc (core vs. shell) are themselves heterogeneous structures with potentially distinct functional roles (45, 46). Future studies using more refined approaches will be needed to resolve how these subdivisions contribute to the pathway- and state-specific activity patterns identified in this work.

In summary, our study demonstrates that mPFC projection pathways make distinct contributions to behaviorally defined affective states in response to environmental cues. Neurons projecting to the BLA preferentially encode negatively valenced states and respond strongly to anxiogenic contexts, whereas those projecting to the NAc are more engaged during positively valenced states, including exploration and social preference. These findings advance our understanding of the neural circuitry underlying emotional processing and behavioral adaptation, and they provide a framework for investigating how disruptions in these pathways may contribute to neuropsychiatric disorders. Importantly, while our imaging results reveal distinct activity patterns in the mPFC→BLA and mPFC→NAc pathways during emotionally salient behaviors, they are correlational in nature. Establishing a causal role will require pathway- or ensemble-specific manipulations during behavior, an essential next step for directly linking neural dynamics to affective regulation.

Methods and materials

Animal care and usage followed NIH Guidelines and received approval from the Institutional Animal Care and Use Committee of George Washington University and the University of Tennessee Health Science Center Laboratory Animal Care Unit.

Experimental Animals

CaMKII-Cre mice with a pure C57BL/6 background were acquired from Jackson Lab (JAX#005359). For excitatory neuron imaging experiments, female CaMKII-Cre mice were bred by mating male CaMKII-Cre mice with female mice from the 129S1/SvlmJ strain (JAX#002448). These experimental CaMKII-Cre mice were used in accordance with the same procedural guidelines. Surgeries were performed on mice at approximately 4 months of age, adhering to the established experimental protocol. Animals were housed in groups of 4 to 5 per cage in a controlled environment with a temperature of 23 ± 1°C, humidity set at 50 ± 10%, and a 12-hour light-dark cycle. Standard mouse chow and water were provided ad libitum.

Virus injection and gradient-index (GRIN) lens implantation

In the context of imaging excitatory neurons, we employed the AAV1-EF1a-flex-GCaMP6m virus acquired from Baylor College of Medicine. Following established protocols (21), CaMKII - Cre mice were anesthetized and securely positioned within a Neurostar stereotaxic frame from Tübingen, Germany. Employing a high-speed rotary stereotaxic drill (Model 1474, AgnTho’s AB, Lidingö, Sweden), we conducted a unilateral injection of the retrograde virus into either the left NAc region (anterior-posterior AP: +1.34 mm, medial-lateral ML: -1.1 mm, dorsal-ventral DV: 4.4mm) or the Amygdala region (AP: -1.6 mm, M/L: -2.9 mm, DV: 4.5mm) using the Nanojector II system from Drummond Scientific. The injection comprised 300 nL of virus diluted with 300 nL of phosphate buffer solution (PBS), administered at a controlled rate of 30 nL/min. Post-injection, the needle was retained in place for an additional 5 minutes to optimize virus diffusion.

Subsequent to viral injection, a precise 1.1 mm-diameter craniotomy was conducted at AP: +1.95 mm, M/L: -0.5 mm coordinates. These values were determined relative to bregma: +1.95 mm AP, -0.35 mm ML, -2.3 to -2.5 DV, employing a high-resolution atlas. Following this, a 1-mm diameter GRIN lens (Inscopix, Palo Alto, CA) was progressively lowered into the left PL region (AP: +1.95 mm; ML: ± 0.35 mm; DV: -2.1∼-2.3 mm) at a rate of 50 μm/min, situated 0.2 mm above the virus injection site, and then cemented in place using Metabond S380 (Parkell). Subsequently, the mice were allowed to recover on a heating pad and were closely monitored over the ensuing 7 days, during which they received analgesic treatment.

Baseplate attachment

Around three to four weeks post-surgery, we assessed virus expression in the anesthetized mice using a miniaturized microscope from Inscopix, based in Palo Alto, CA. Once GCaMP+ neurons were clearly visible, we proceeded to attach the microscope with a baseplate onto the mouse’s skull window. This setup was then gradually lowered to determine the optimal focus plane. Subsequently, the baseplate was securely affixed to the skull using dental cement and covered with a protective cap, while the microscope remained unattached. Ahead of the behavioral tests, the mice were familiarized with the test room environment, during which a dummy microscope was mounted and handled for approximately 5 to 7 days, with sessions lasting 30 to 40 minutes each day.

Selection of animals

Mice were chosen based on specific criteria: for the observation of Ca2+ signals in the mPFC, mice were excluded post-hoc if (1) the GRIN lens was positioned outside of the prelimbic cortex, (2) GCamp6 was not expressed within prelimbic areas, (3) significant virus expression was observed outside of the prelimbic region, and de(4) the imaging plane was obstructed by blood or debris. In the identification of mPFC neural ensembles, mice exhibiting no center entry behavior were excluded, as the classification was grounded in the neural responses to transitions in location.

Determination of sample size

The sample sizes for behavioral experiments were established following the prevailing standards in behavioral neuroscience research for mice. This approach considers the minimum number of mice needed to detect statistical significance, with an α level of 0.05, and a statistical power of 80% or higher. Our post-hoc analysis demonstrates a statistical power of 85%.

In terms of neuron count, we recorded 11-83 neurons expressing GCaMP6m for the mPFC→BLA group (with an average of 38) and 18-54 for the mPFC→NAc group (with an average of 39). This range was influenced by various factors, including the injected volume of viral GCaMP6m, expression levels, and the efficiency of Cre-mediated recombination. Following the data processing phase, which involved the identification of neurons through principal component and independent component analyses (PCA-ICA), approximately 1-9% of the identified components were identified as artifacts and subsequently excluded. The remaining components were classified as neurons and employed for subsequent analyses.

Behavioral tests

The behavioral tests spanned three consecutive days, maintaining consistent schedules throughout. To ensure cleanliness, the chamber was meticulously cleaned with 70% ethanol between trials. The Topscan behavior analysis system from Clever Sys, VA, was employed to monitor the animals’ behaviors. This system also sent a TTL signal at the start of each test to trigger the microscope recording of neuronal activity.

Modified Three-Chamber (mTC) Test

The mTC test adhered to a previously established protocol (47) with several enhancements. The conventional three-chamber apparatus was transformed into a single open box (45 x 10 x 20 cm) featuring two small removable lateral chambers (10 x 10 x 40 cm). These chambers were separated by thin, spaced metal wires, permitting mice to interact with stimuli. Before the test, the subject mouse underwent a 10-minute habituation period within the open box devoid of stimuli. During testing, an age- and weight-matched, unfamiliar same-sex conspecific (referred to as the first social stimulus, M1) and an inanimate object (non-social stimulus, O) were randomly placed into the two lateral chambers. Subsequently, the subject mouse was positioned in the center of the open box and allowed to explore freely during a 10-minute testing session. Several behavioral parameters were assessed, including time spent in social interaction, object interaction, the social zone (SZ), the object zone (OZ), the transition zone, and grooming behavior.

Open field test (OFT)

The test was modified from a previous report (14). It was conducted in a square box (dimensions: 50 × 50 × 50 cm). The mouse was gently placed in the central field and allowed to explore freely during a 10-minute testing session. Locomotor activity was recorded by a camera. The center is defined as the central 25 cm x 25 cm square area, while the corner is a sector area with a 12.5 cm radius in each corner. Total distance traveled and time spent in each area, sniffing, and grooming behaviors were analyzed using the TopScan automated behavioral analysis system (CleverSys Inc., Reston, VA). “Sniffing” behavior was defined as episodes when the mouse’s nose was within approximately 1 cm of the wall and moving in small scanning motions without locomotion away from the wall, lasting for at least 1 second. To ensure clear separation from “corner” behavior, sniffing events occurring within the defined corner zones were excluded from the corner category and analyzed separately. As a result, corner time does not include sniffing episodes that occurred in corners.

Elevated plus maze (EPM) test

It was conducted as previously described (16). The EPM apparatus, 40 cm high from the floor, consists of two open arms (35 × 10 cm) and two closed arms (35 × 10 cm), those two parts stretch perpendicular to each other and connect to a center platform (5 cm). Mice were placed in the center zone facing an open arm and allowed to explore the maze freely for 10 min. The time spent in open arms, closed arms, and behaviors of head dipping, sniffing, and grooming were analyzed.

Dominance tube test

The tube test protocol was modified from Wang et al (23). In brief, a clear acrylic tube with a 30 cm length and 2.8 cm inside diameter in which one mouse can pass or backward through the tube fluently, but cannot make a U-turn or climb through another mouse. The mice underwent a three-day training in which all mice went through the tube ten times per day, five times from each side, without another competitor in the tube. The training procedure allows the mice to adapt to the test procedure and environment. The mice, in a very rare case, cannot be trained and were excluded from further testing. In the testing procedure, tests were conducted in a pair-wise style, and the number of times won by individual animals was recorded to determine the hierarchical ranking. The mouse that forced its competitor out of the tube was declared the “winner”, or dominant in this situation. The mouse that was retreated is then designated the “loser” or subordinate. In most cases, the competition was completed within 2 minutes, or the tests were repeated. The rank is considered stable without a ranking change in all mice for at least five consecutive days. Groups were categorized as “non-stable ranking” if they did not form a stable ranking within 2 weeks. We single-housed the mice for 5 days and subjected them to a round-robin tube test tournament (three trials per day for 7 trials) to determine social ranking. Twenty-four hours following the last tube test, the mice were subjected to other behavioral tests or in vitro studies, as described in the results.

Behavior data analysis

Earlier version: Behavioral data was meticulously tracked through aerial videography from an overhead perspective utilizing the Topscan behavioral data acquisition software (CleverSys, Reston, VA). This software allowed for the precise tracking and definition of the 2D locations of mice in various areas, including the OF center, corner area, EPM open-P, open-NP, closed-P, closed-NP arms, and mTC social zone, object zone, and middle zone. Additionally, specific behaviors such as sniffing, grooming, head-dipping, and social/object interaction were recognized and quantified by the software.

Revised version: Behavioral data was meticulously tracked through aerial videography from an overhead perspective utilizing the Topscan behavioral data acquisition software (CleverSys, Reston, VA). This software allows for the precise tracking and definition of the 2D locations of mice in various areas, including the OF center, corner area, EPM open-P, open-NP, closed-P, closed-NP arms, and mTC social zone, object zone, and middle zone. Additionally, specific behaviors such as sniffing, grooming, head-dipping, and social/object interaction were recognized and quantified by the software.

Ca2+ imaging with Miniature microscope

Imaging of freely moving mice was conducted using a head-mounted miniaturized microscope (nVista HD 2.0.4, Inscopix, Palo Alto, CA), as depicted in Fig. 1A. This microscope was synchronized with the Topscan system through a TTL pulse, enabling simultaneous acquisition of Ca2+ signals and behavioral video. Prior to imaging, the microscope was securely affixed to the mouse’s head. The imaging data were obtained at a frame rate of 15 Hz with a resolution of 1024 x 1024 pixels. LED power settings ranged from 0.3 to 1 mW, while gain settings were adjusted to 1 to 2 based on fluorescence intensity. Importantly, each individual mouse utilized the same imaging parameters consistently across all three experimental sessions.

Histology

Recording sites were meticulously validated through histological examination of lesions created during the lens implantation procedure. Mice were anesthetized via intraperitoneal injection, employing a combination of ketamine (400 mg/kg) and xylazine (20 mg/kg). Following this, transcardial perfusion was conducted using PBS, followed by 4% paraformaldehyde (PFA). The brains of the mice, complete with their skulls and baseplates, were post-fixed with 4% PFA for 3 days. Subsequently, the brains were extracted and sectioned into slices measuring 50-100 µm using a vibrating slicer (Vibratome Series 1000, St. Louis, MO). These sections were then mounted onto slides. To label cell nuclei, slides were subjected to incubation and storage in a 1:1000 Hoechst solution in 1x PBS (Invitrogen, Carlsbad, CA). Brain slides were meticulously imaged to precisely determine the placement of the GRIN lens and the extent of viral expression. This imaging process was conducted employing a Confocal Microscope (Zeiss LSM 710, Oberkochen, Germany).

Ca2+ image processing

Ca2+ images were processed offline using Inscopix Data Processing software (version 1.3.1). In brief, frames collected within a single day were concatenated into a stack and subsequently subjected to preprocessing, spatial filtering, and motion correction. To normalize the Ca2+ signal, the average projection of the filtered video was established as the background fluorescence (F0). Instantaneous normalized Ca2+ fluorescent signals (ΔF/F) was calculated according to the formula, (ΔF/F)i=(Fi-F0)/F0, where i represents each frame. Individual cells were then identified using the principal component and independent component (PCA-ICA) analyses with no spatial or temporal down-sampling. Regions of interest (ROI) were selected based on signal and image criteria, with any components that did not correspond to single neurons being discarded.

Earlier version: Time-stamped traces of neurons were exported into data files formatted for custom-written Python scripts used for subsequent analysis. Ca2+ transients (events) were identified for each cell using a peak-finding algorithm, and frequency (transient rate) and amplitude (ΔF/F) data were processed using custom-written scripts. When analyzing frequency or amplitude changes in response to mouse behavior or location, frames of image and behavioral data were aligned and marked with corresponding behavioral event labels. For neuronal activity normalization, we employed the min-max normalization method to scale values between 0 and 1. In this method, the maximum neuron activity value was transformed into 1, while all other values were scaled to decimals between 0 and 1 using the equation: , where x = (x1,x2,…,xn) represents an array of neuron frequency or amplitude values, and zi denotes the normalized value.

Revised version: Time-stamped traces of neurons were exported into data files formatted for custom-written Python scripts used for subsequent analysis. Ca2+ transients (events) were identified for each cell using a peak-finding algorithm, and frequency (transient rate) and amplitude (ΔF/F) data were processed using custom-written scripts. When analyzing frequency or amplitude changes in response to mouse behavior or location, frames of image and behavioral data were aligned and marked with corresponding behavioral event labels. For neuronal activity normalization, we employed the min-max normalization method to scale values between 0 and 1. In this method, the maximum neuron activity value was transformed into 1, while all other values were scaled to decimals between 0 and 1 using the equation: zi=xi−min(x)max(x)−min(x), where x = (x1,x2,…,xn) represents an array of neuron frequency or amplitude values, and zi denotes the normalized value.

Identification of location-modulated neural ensembles

To identify location-modulated neural ensembles in each session, we assessed each neuron’s response preference to a specific location exploration. Initially, we computed the actual similarity (Sa) between the Ca2+ trace vectors (ck) and behavior events (b), using the formula: 2b ⋅ ck/(|b|2 + |ck|2) (18). Then, the behavior vector was subjected to random shuffling to compute a new similarity (S) with a neural trace for a given neuron. This shuffling process was repeated 5000 times, generating an S distribution histogram. Neurons were classified as ON neurons if their Sa values exceeded the 99.95th percentile of the S distribution.

Following neuron identification, we calculated the proportions of different location-modulated ON neurons for each mouse. Simultaneously, the transient rate and amplitude of these neuron ensembles were computed for each mouse in each paradigm.

Neuron alignment between tests

Earlier version: To identify and compare the same neurons across tests conducted on different days, thus enabling cross-test comparisons, we implemented a global alignment procedure. Initially, we reconstructed neuron distribution images based on their pixel locations and shapes. Subsequently, we converted each neuron image into a vector format and assessed its similarity using cosine similarity, as indicated by the formula below. The two parameters, degree of rotation (A) and pixel shift (B), were iteratively adjusted to maximize the cosine similarity of image vectors between tests. Given that the images from all three tests were collected under identical settings, the scale for the neuron images remained constant and did not require adjustment. Through this multi-start method, we achieved the global optimal solution for the parameters. Ultimately, we determined the number of overlapping neuron images across tests, which was then used for subsequent comparative analyses.

Revised version: To identify and compare the same neurons across tests conducted on different days, thus enabling cross-test comparisons, we implemented a global alignment procedure. Initially, we reconstructed neuron distribution images based on their pixel locations and shapes. Subsequently, we converted each neuron image into a vector format and assessed its similarity using cosine similarity, as indicated by the formula below. The two parameters, degree of rotation (A) and pixel shift (B), were iteratively adjusted to maximize the cosine similarity of image vectors between tests. Given that the images from all three tests were collected under identical settings, the scale for the neuron images remained constant and did not require adjustment. Through this multi-start method, we achieved the global optimal solution for the parameters. Ultimately, we determined the number of overlapping neuron images across tests, which was then used for subsequent comparative analyses. Similarity=cos⁡(θ)=A⋅B‖A‖‖B‖

Principal Component Analysis (PCA)

Population neural activity was analyzed using principal component analysis (PCA) to reduce dimensionality and visualize population dynamics across behavioral states. Calcium event rates were z-scored for each neuron across time before concatenation across mice to generate a population activity matrix (frames × neurons) for each pathway. PCA was applied to this pooled matrix using the covariance method implemented in MATLAB (MathWorks) and scikit-learn (Python), centering each neuron’s activity by its mean to project neurons from different animals into a common coordinate space.

The first three principal components (PCs) were used to compute neural trajectories and distance metrics, as they captured the majority of the variance in population activity (∼49% for all neurons and ∼59% for center-ON neurons in the mPFC→BLA pathway). To quantify neural representation differences across behaviors, we calculated the Euclidean distance from each behavioral cluster to the Corner cluster within the same PC space. Corner behavior was chosen as the reference because mice spent most of their time in the corners during the open-field test, representing a stable low-anxiety baseline state. Euclidean distances between behavioral clusters were computed using the centroid coordinates of each behavior (Center, Sniffing, Grooming, and Corner) in the two-dimensional PCA space defined by the first two principal components (PC1 and PC2), averaged across mice for each pathway.

Cross-animal analysis

For cross-animal population analyses (e.g., Fig. 1I), neuronal calcium activity from all mice was concatenated along the neuron dimension to create a population matrix (frames × neurons). PCA was applied directly to this pooled dataset. Because PCA centers each neuron’s activity by its mean, neurons from different animals were projected into a shared coordinate frame without requiring further normalization. This approach enabled consistent visualization of behavioral state trajectories across animals in a common PC space.

Optogenetics and patch-clamp recording

Viral injection for optogenetics

We stereotaxically injected an AAV2-CaMKIIα-hChR2(H134R)-EYFP (UNC virus core, 0.5 µl of 1.00 E13 viral genomes/ml) into mPFC bilaterally (relative to bregma: +1.9 mm AP; ±0.3 mm ML; -2.0 mm DV) using a Hamilton microsyringe (Model 1701N) and a WPI microsyringe pump mounted to a KOPF Model 940 Small Animal Stereotaxic Instrument.

After viral injection, mice were recovered for 2 weeks to allow for the AAVs’ efficacy. We then performed anterograde tracing from the mPFC to the NAc and BLA in brain slices. After euthanasia using 2% isoflurane, the mice underwent transcardiac perfusion with 4% PFA to fix the entire brain. After a 6-hour fixation period, 50 µm coronal slices, encompassing the mPFC, NAc, and BLA, were dissected using a vibratome (Leica VT-1000S) and collected in ice-cold PBS. These sections were transferred onto gelatin-coated slides (Thermo Fisher Scientific) and mounted onto cover slides (Electron Microscopy Sciences). Nuclei were labeled using the Gold antifade reagent with DAPI (Invitrogen). Fluorescence images were captured and visualized at 10x and 20x magnification under a Leica microscope (DMRXA2, Leica Microsystems) equipped with a digital camera and imaging software.

Brain slice preparation

The patch clamp test was conducted following the completion of the tube test. Mice were anesthetized with 2% isoflurane and subsequently decapitated. The brain skull was swiftly removed and immersed in an ice-cold cutting solution composed of the following concentrations (in mM): 205 sucrose, 5 KCl, 1.25 NaH2PO4, 5 MgSO4, 26 NaHCO3, 1 CaCl2, and 25 glucose; with an osmolarity of 300-320 mOsm. The brain was then sliced into 300 µm coronal sections using a vibratome (Leica VT-1000S), encompassing the mPFC, NAc, and BLA.

Post-sectioning, slices were maintained in a normal artificial cerebrospinal fluid (ACSF) solution for whole-cell recording, containing the following concentrations (in mM): 115 NaCl, 2.5 KCl, 2 CaCl2, 1 MgCl2, 1.25 NaH2PO4, 17 glucose, 25 NaHCO3; with an osmolarity of 280-300 mOsm. The slices were bubbled with 95% O2 / 5% CO2 for one hour at ambient temperature (22 ± 3°C). Whole-cell patch-clamp recordings were conducted by placing the slices in a submerged recording chamber and perfusing them with ACSF at a rate of 1 ml/min.

Earlier version: Patch-clamp recording

Revised version: Optogenetic stimulation and Patch-clamp recording

The pipette solution for EPSC recordings contained (in mM): 145 Cs-methanesulfonate (Cs-SO3CH3), 8 NaCl, 1 MgCl2, 0.38 CaCl2, 1 Cs-EGTA, and 10 HEPES (mOsm 280 - 300, adjusted to pH 7.25 with CsOH). For recordings in the current-clamp mode, CsSO3CH3 was replaced by K-SO3CH3. The pipette resistance was 4-5 MΩ, and the series resistance was 8-10 MΩ. No series resistance compensation was applied. Brain slices were recorded through an Axon Instruments Multiclamp 700B amplifier, sampled at 10 kHz, and digitized with a Digidata 1440A (Molecular Devices software). The cell morphology was identified under an Olympus BX51 fluorescence microscope. 100 µM picrotoxin was included in the ACSF throughout the recording.

Earlier version: A blue light pulse (PlexBright LED, 470nm) was utilized to activate mPFC fibers expressing ChR2. To validate the expression of ChR2 in the mPFC, inward current in voltage-clamp recording was assessed at –80 mV for 1 second, with a light intensity of 1 µW. Furthermore, voltage spikes were examined in current clamp mode at frequencies of 5, 10, and 20 Hz, induced by a 1-ms light pulse with an intensity of 1 µW.

Revised version: Light-evoked responses were elicited by positioning a fiber-optic directly into the recorded slice to activate ChR2-expressing terminals. A blue light pulse (PlexBright LED, 470nm) was utilized to activate mPFC fibers expressing ChR2. To validate the expression of ChR2 in the mPFC, inward current in voltage-clamp recording was assessed at -80 mV for 1 second, with a light intensity of 1 µW. Furthermore, light-evoked EPSCs were examined at frequencies of 5, 10, and 20 Hz, induced by a 1-ms light pulse with an intensity of 1 µW.

Added in the revised version: To elicit synaptic responses in the NAc and BLA through light pulses targeting mPFC fibers, slices were illuminated at 30-second intervals, with a short duration of 1-5 ms for NAc neurons and 1-10 ms for BLA neurons. Because opsin expression and light delivery can vary across animals and slices, light-evoked responses were normalized within each slice to account for inter-slice variability. LED pulse duration (1-10 ms) was adjusted for each slice to reliably elicit responses, and recordings were performed at consistent mPFC→NAc or mPFC→BLA locations. These parameters were optimized to evoke robust responses, but do not imply identical absolute stimulation intensity or synaptic strength across experiments. EPSCs were induced by incrementally increasing light intensity (1, 2, 5, and 10 µW) at a holding potential of -80 mV, normalized by dividing by the response at 10 µW. PPR responses at 50 and 75 ms intervals were measured at an intensity of 1 µW, and the amplitude of the second response was divided by that of the first response.

Earlier version: To elicit synaptic responses in the NAc and BLA through light pulses targeting mPFC fibers, slices were illuminated at 30-second intervals, with a short duration of 1-5 ms for NAc neurons and 1-10 ms for BLA neurons. EPSCs were induced by incrementally increasing light intensity (1, 2, 5, and 10 µW) at a holding potential of –80 mV, normalized by dividing by the response at 10 µW. PPR responses at 50 and 75 ms intervals were measured at an intensity of 1 µW, and the amplitude of the second response was divided by that of the first response. To ensure robust and reproducible synaptic responses, we titrated pulse durations within these ranges during pilot experiments to identify the minimal duration required to elicit stable EPSCs without inducing depolarization block or excessive synaptic depression, as described in optogenetic studies (48, 49). These ranges (1-5 ms for NAc, 1-10 ms for BLA) were selected to balance activation efficiency with the prevention of artifacts, such as those caused by prolonged light exposure. Our experimental design accounted for this by normalizing EPSC amplitudes to the response at a standard light intensity (10 µW) across all experiments. This normalization minimizes the impact of pulse duration variability on the steepness of the input-output curves. Additionally, PPR measurements were conducted at a fixed intensity (1 µW) and consistent interstimulus intervals, ensuring that comparisons of synaptic properties were not confounded by pulse duration differences. Electrophysiological recordings were performed in acute coronal slices prepared from behaviorally naïve mice following the tube test to assess pathway-specific synaptic properties. Viral injections were targeted to the mPFC using standardized stereotaxic coordinates to ensure consistent labeling of projections to the BLA or NAc. During recordings, cells were selected from anatomically defined subregions (BLA or NAc shell) identified under IR-DIC optics based on established landmarks, such as the external capsule and anterior commissure, to maintain spatial consistency across experiments.

Revised version: To ensure robust and reproducible synaptic responses, we titrated pulse durations within these ranges during pilot experiments to identify the minimal duration required to elicit stable EPSCs without inducing depolarization block or excessive synaptic depression, as described in optogenetic studies (48, 49). These ranges (1-5 ms for NAc, 1-10 ms for BLA) were selected to balance activation efficiency with the prevention of artifacts, such as those caused by prolonged light exposure. Our experimental design accounted for this by normalizing EPSC amplitudes to the response at a standard light intensity (10 µW) across all experiments. This normalization minimizes the impact of pulse duration variability on the steepness of the input-output curves. Additionally, PPR measurements were conducted at a fixed intensity (1 µW) and consistent interstimulus intervals, ensuring that comparisons of synaptic properties were not confounded by pulse duration differences. Electrophysiological recordings were performed in acute coronal slices prepared from behaviorally naïve mice following the tube test to assess pathway-specific synaptic properties. Viral injections were targeted to the mPFC using standardized stereotaxic coordinates to ensure consistent labeling of projections to the BLA or NAc. During recordings, cells were selected from anatomically defined subregions (BLA or NAc shell) identified under IR-DIC optics based on established landmarks, such as the external capsule and anterior commissure, to maintain spatial consistency across experiments.

To determine the AMPA/NMDA current ratio, we measured the peak amplitude of AMPA-EPSCs at -70 mV and the peak amplitude of NMDA-EPSCs at +40 mV at 70 ms after onset (50) with an intensity of 2 µW.

Measurement of corticosterone levels in mPFC tissue

Earlier version: Corticosterone concentrations in the mPFC were quantified using a commercially available Cortisol ELISA kit (IBL America, IB79175) following the manufacturer’s instructions, with minor modifications for tissue samples. Briefly, mPFC tissues were rapidly dissected on ice and weighed using a micro-scale. 500μl of Ethyl Acetate was added onto the samples prior to homogenization on ice with a hand-held homogenizer: 10 seconds on and 10 seconds off 4 times. An additional 500μl of Ethyl Acetate was added onto the samples and then thoroughly vortexed (∼2min) and centrifuged at 500 g for 5 min at 4°C and left on dry ice for at least 6 hours. Supernatant was transferred to a new 2 ml Eppendorf tube and dried using a SpeedyVac rotary vacuum evaporator at 50°C. Samples were then stored at –80°C until used. Before performing the assay, samples were resuspended in 25 μl of 100% Ethanol. Resuspension of the samples in 0 calibrator or PBS resulted in the need for a large volume, reducing the concentration of corticosterone below the detection limit. The ELISA assay was performed following the manufacturer’s instructions, except that the assay buffer was added first, and 10μl of the samples were added on top of that. This is to avoid direct exposure of the antibodies to ethanol. Absorbance was then measured at 450 nm using an EnVision plate reader. The concentrations were calculated from a standard curve and normalized to total protein content determined by the BCA assay.

Revised version: Corticosterone concentrations in the mPFC were quantified using a commercially available Cortisol ELISA kit (IBL America, IB79175) following the manufacturer’s instructions, with minor modifications for tissue samples. Briefly, mPFC tissues were rapidly dissected on ice and weighed using a micro-scale. 500μl of Ethyl Acetate was added onto the samples prior to homogenization on ice with a hand-held homogenizer: 10 seconds on and 10 seconds off 4 times. An additional 500μl of Ethyl Acetate was added onto the samples and then thoroughly vortexed (∼2min) and centrifuged at 500 g for 5 min at 4°C and left on dry ice for at least 6 hours.

Added in the revised version: Supernatant was transferred to a new 2 ml Eppendorf tube and dried using a SpeedyVac rotary vacuum evaporator at 50°C. Samples were then stored at -80°C until used. Before performing the assay, samples were resuspended in 25 μl of 100% Ethanol. Resuspension of the samples in 0 calibrator or PBS resulted in the need for a large volume, reducing the concentration of corticosterone below the detection limit. The ELISA assay was performed following the manufacturer’s instructions, except that the assay buffer was added first, and 10μl of the samples were added on top of that. This is to avoid direct exposure of the antibodies to ethanol. Absorbance was then measured at 450 nm using an EnVision plate reader. The concentrations were calculated from a standard curve and normalized to total protein content determined by the BCA assay.

Statistics

Earlier version: All statistical analyses were performed using SPSS (version 24, IBM, Armonk, NY), Microsoft Excel (Redmond, WA), and custom Python scripts. Data were first tested for normality using the D’Agostino–Pearson test. For comparisons between two dependent conditions (e.g., data obtained from the same mouse), a two-tailed paired t-test was applied; for independent two-group comparisons, an unpaired two-tailed t-test was used. When comparing more than two groups, one-way ANOVA followed by Tukey’s post hoc multiple-comparison test was conducted. For datasets involving two experimental factors, two-way mixed-effects ANOVA with Bonferroni-corrected post hoc tests was performed. Statistical significance was defined as p < 0.05 (p < 0.01, highly significant), and non-significant results (e.g., p > 0.05 in paired t-tests) are reported where relevant. All data are presented as mean ± SEM, unless otherwise indicated.

Revised version: All statistical analyses were performed using SPSS (version 24, IBM, Armonk, NY), Microsoft Excel (Redmond, WA), and custom Python scripts. Data were first tested for normality using the D’Agostino-Pearson test. For comparisons between two dependent conditions (e.g., data obtained from the same mouse), a two-tailed paired t-test was applied; for independent two-group comparisons, an unpaired two-tailed t-test was used. When comparing more than two groups, one-way ANOVA followed by Tukey’s post hoc multiple-comparison test was conducted. For datasets involving two experimental factors, two-way mixed-effects ANOVA with Bonferroni-corrected post hoc tests was performed. Statistical significance was defined as p < 0.05 (p < 0.01, highly significant), and non-significant results (e.g., p > 0.05 in paired t-tests) are reported where relevant. Normality was assessed using the Shapiro-Wilk test. When normality assumptions were violated or sample sizes were small (n < 15), non-parametric Mann-Whitney U tests were additionally performed. Effect sizes were reported as Cohen’s d for parametric tests and rank-biserial correlation for non-parametric tests. All data are presented as mean ± SEM, unless otherwise indicated.

To validate the significance of distance metrics in neural state space, permutation-based null controls were implemented by (i) shuffling behavioral labels within each mouse and session while preserving bout counts and trajectory lengths, and (ii) applying circular time-shifts per neuron within bouts to maintain each neuron’s autocorrelation. Observed cluster distances and pathway differences were compared against 10,000 permutation-derived null distributions, with significance defined as exceeding the 95th percentile (p_perm < 0.05).